# Lab 2

1. 提交作業之前，建議可以先點選上方工具列的**Kernel**，再選擇**Restart & Run All**，檢查一下是否程式跑起來都沒有問題，最後記得儲存。
2. 請先填上下方的姓名(name)及學號(stduent_id)再開始作答，例如：
```python
name = "我的名字"
student_id= "B06201000"
```
3. 四個求根演算法的實作可以參考[lab-2](https://yuanyuyuan.github.io/itcm/lab-2.html)，裡面有教學影片也有範例程式可以套用。
4. **Deadline: 10/9(Wed.)**

```python
name = ""
student_id = ""
```

---

# Exercise 2

---

## Kepler's equation

### In celestial mechanics, _**Kepler’s equation**_
### $$M = E - e\,\sin(E)$$
### relates the mean anomaly $M$ to the eccentric anomaly $E$ of an elliptical orbit of eccentricity $e$, where $0 < e < 1$, see [ _Wiki website_ ](https://en.wikipedia.org/wiki/Kepler's_laws_of_planetary_motion) for the details.

---

### 1. Prove that fixed-point iteration using the iteration function 
### $$g(E) = M + e\,\sin(E)$$
### is convergent locally.
\[Hint: You may use Ostrowski’s Theorem mentioned in the lecture note.\]

### proof.

<!-- grade_id: prove_convergent_locally, points: 3, grade: true, solution: true -->
請點此cell兩下開始作答（如要打文字記得選Markdown, 寫程式則選Code, 一個cell不夠可以再新增在下方）

---


### 2. Use the fixed-point iteration scheme in "Q.1" to solve Kepler’s equation for the eccentric anomaly $E$ corresponding to a mean anomaly $M = \dfrac{2\pi}{3}$ and an eccentricity $e = 0.5$.

---

### Part 0. Import libraries

```python
import matplotlib.pyplot as plt
import numpy as np
```

### Part 1. Define the fixed point function

<!-- grade_id: fixed_point_function, points: 0, grade: false, solution: true -->
```python
def fixed_point(
    func,
    x_0,
    tolerance=1e-7,
    max_iterations=5,
    report_history=False,
):
    '''Approximate solution of f(x)=0 on interval [a,b] by the secant method.

    Parameters
    ----------
    func : function
        The target function.
    x_0 : float
        Initial guess point for a solution f(x)=0.
    tolerance: float
        One of the termination conditions. Error tolerance.
    max_iterations : (positive) integer
        One of the termination conditions. The amount of iterations allowed.
    report_history: bool
        Whether to return history.

    Returns
    -------
    solution : float
        Approximation of the root.
    history: dict
        Return history of the solving process if report_history is True.
    '''
    
    # 請參考 hands-on 的 fixed point method
    #
    # ===== 請實做程式 =====
    
    # ====================
    
```

Test your implementaion with the assertion below.

<!-- grade_id: test_fixed_method, points: 1, grade: true, solution: false -->
```python
root = fixed_point(lambda x: x - (x**2 - 4*x + 3.5), 2, tolerance=1e-7, max_iterations=100, report_history=False)

error = np.inf
for solution in np.roots([1, -4, 3.5]):
    if abs(root - solution) < error:
        exact_solution = solution
        error = abs(root - solution)

assert error < 1e-7
```

### Part 2. Assign values to variables anomaly mean "_M_" and eccentricity  "_e_".

### $$M = \frac{2\pi}{3}\,\,\,\,\,\text{and}\,\,\,\,\,e = 0.5$$

<!-- grade_id: assign_values, points: 0, grade: false, solution: true -->
```python
# Hint:
# M = ?
# e =?
# ===== 請實做程式 =====

# ====================
```

<!-- grade_id: M_and_e, points: 1, grade: true, solution: false -->
```python
print('M =', M)
print('e =', e)

```

### Part 3. Define the function of Kepler’s equation
<br>

#### _**Recall Kepler’s equation**_ :
#### $$M = E - e\,\sin(E).$$
#### So we let the function $f(E) = E - e\,\sin(E) - M$, then
#### $$g(E) = E - f(E) = M + e\,\sin(E)$$

#### For the instance:

If we want to implement "$\sin(x)$", we will call `np.sin(x)` with numpy in python.


<!-- grade_id: kepler_equation, points: 0, grade: false, solution: true -->
```python
def f(E):
    # Hint: return ...
    # ===== 請實做程式 =====
    
    # ====================

def g(E):
    # Hint: return ...
    # ===== 請實做程式 =====
    
    # ====================
```

<!-- grade_id: test_f_and_g, points: 2, grade: true, solution: false -->
```python
print('M =', M)

# f(0) = -M, g(0) = M
print('f(0) =', f(0))
print('g(0) =', g(0))

```

### Part 4. Plot the function $f(E)$ and $g(E)$

```python
fig, ax = plt.subplots(figsize=(16, 4))
search_range = np.arange(-10, 10, 0.01)
ax.plot(search_range, f(search_range))
ax.set_title(r'$f(E) = E - e\,\sin(E) - M$')
ax.grid(True)
ax.axhline(y=0)
ax.axvline(x=0)
plt.show()
```

```python
fig, ax = plt.subplots(figsize=(16, 4))
search_range = np.arange(-10, 10, 0.01)
ax.plot(search_range, g(search_range))
ax.set_title(r'$g(E) = M + e\,\sin(E)$')
ax.grid(True)
ax.axhline(y=0)
plt.ylim(1,3)
plt.show()
```

### Part 5. Find the solution of "$E$"

<!-- grade_id: solution_of_E, points: 0, grade: false, solution: true -->
```python
root = fixed_point(
    # ===== 請實做程式 =====
    
    # ====================
)
```

<!-- grade_id: the_root_of_E, points: 3, grade: true, solution: false -->
```python
print('My estimation of root:', root)

```

---

### 3. An “ _exact_ ” formula for $E$ is known: $$E = M + 2 \sum_{m=1}^\infty \frac{1}{m} J_m(me)\,\sin(mM);$$ where $J_m(x)$ is the Bessel function of the first kind of order $m$. 

### Use this formula to compute $E$. How many terms are needed to produce the value obtained in "Q.2" until convergence?

### Part 0. Import package

```python
from scipy.special import jn  # Bessel function
```

### Part 1. Define the function

For the convenience, we define the function $h(m)$ as
$$h(m) \triangleq \frac{2}{m} J_m(me)\,\sin(mM)$$

If we want to implement " _**Bessel function**_ " $J_m(x)$, we can call `jn(m,x)` in Python.

<!-- grade_id: function_h, points: 0, grade: false, solution: true -->
```python
def h(m):
    # Hint: return ...
    # ===== 請實做程式 =====
    
    # ====================
```

<!-- grade_id: h, points: 1, grade: true, solution: false -->
```python
# test the function of h
print('h(1) =', h(1))
assert round(h(1), 5) == 0.41962

```

### Part 2. Find how many terms we need to achieve the result obtained Q.2 in a tolerance $10^{-7}$.

That is to find _num_terms_ such that

$$ \left|\ \text{root} - \left(M + \sum_{k=1}^{\text{num_terms}} h(k)\right)\ \right| < 10^{-7} $$

For example, the following cell shows the implmentation with only 1 term. 

```python
LHS = root
RHS = M + h(1)
error = abs(LHS-RHS)
print('Left hand side is the estimation of root by the fixed-point method:', LHS)
print('Right hand side is the approximation by the formula in only 1 term:', RHS)
print('The error between LHS and RHS:', error)
```

<!-- grade_id: find_the_number_of_term, points: 0, grade: false, solution: true -->
```python
LHS = root
RHS = M + h(1)
num_terms = 1
tolerance = 1e-7

# ===== 請實做程式 =====

# ====================
```

<!-- grade_id: number_of_term, points: 3, grade: true, solution: false -->
```python
print('Number of terms to approximate:', num_terms)

```