# Lab 2

1. 提交作業之前，建議可以先點選上方工具列的**Kernel**，再選擇**Restart & Run All**，檢查一下是否程式跑起來都沒有問題，最後記得儲存。
2. 請先填上下方的姓名(name)及學號(stduent_id)再開始作答，例如：
```python
name = "我的名字"
student_id= "B06201000"
```
3. 四個求根演算法的實作可以參考[lab-2](https://yuanyuyuan.github.io/itcm/lab-2.html)，裡面有教學影片也有範例程式可以套用。
4. **Deadline: 10/9(Wed.)**

```python
name = ""
student_id = ""
```

---

# Exercise 1 - Bisection

### Use the bisection method to find roots of
### $$f(x)=cosh(x)+cos(x)-c \text{, for } c=1, 2, 3,$$

---

### Import libraries

```python
import matplotlib.pyplot as plt
import numpy as np
```

### 1. Define a function $g(c)(x)=f(x)=cosh(x)+cos(x)-c$ with parameter $c=1, 2, 3$.

<!-- grade_id: cell-d3866c0338b6c08a, points: 0, grade: false, solution: true -->
```python
def g(c):
    assert c == 1 or c == 2 or c == 3
    def f(x):
        # Hint: return ...
        # ===== 請實做程式 =====
        
        # ====================
    return f
```

Pass the following assertion.

<!-- grade_id: cell-b59c94b754b1fc9e, points: 1, grade: true, solution: false -->
```python
assert g(1)(0) == np.cosh(0) + np.cos(0) - 1
```

### 2. Implement the algorithm

<!-- grade_id: cell-cf778e6b6ea6b307, points: 0, grade: false, solution: true -->
```python
def bisection(
    func,
    interval,
    max_iterations=5,
    tolerance=1e-7,
    report_history=False,
):
    '''
    Parameters
    ----------
    func : function
        The target function
    interval: list
        The initial interval to search
    max_iterations: int
        One of the termination conditions. The amount of iterations allowed.
    tolerance: float
        One of the termination conditions. Error tolerance.
    report_history: bool
        Whether to return history.

    Returns
    -------
    result: float
        Approximation of the root.
    history: dict
        Return history of the solving process if report_history is True.
    '''

    # ===== 請實做程式 =====
    
    # ====================
```

Test your implementation with the assertion below.

<!-- grade_id: cell-4d88293f2527c82d, points: 1, grade: true, solution: false -->
```python
root = bisection(lambda x: x**2 - x - 1, [1.0, 2.0], max_iterations=100, tolerance=1e-7, report_history=False)
assert abs(root - ((1 + np.sqrt(5)) / 2)) < 1e-7
```

---

### 3. Answer the following questions under the case $c=1$.

### Plot the function to find an interval that contains the zero of $f$ if possible.

<!-- grade_id: cell-b0ab2d19e453ca85, points: 0, grade: false, solution: true -->
```python
c = 1
f = g(c)

# Hint: search_range = np.arange(左端點, 右端點, 點與點之間距), 
# e.g. search_range = np.arange(0.0, 1.0, 0.01)
# ===== 請實做程式 =====

# ====================

fig, ax = plt.subplots(figsize=(9, 4))
ax.plot(search_range, f(search_range))
ax.set_title(r'$f(x)=cosh(x)+cos(x)-c$, $c=$%d' % c)
ax.grid(True)
ax.axhline(y=0)
ax.axvline(x=0)
plt.show()
```

### According to the figure above, estimate the zero of $f$.

#### For example, 

~~~python
root = 3        # 單根
root = -2, 1    # 多根
root = None     # 無解
~~~

<!-- grade_id: cell-04d5f93c79b8260d, points: 0, grade: false, solution: true -->
```python
# Hint: root = ?
# ===== 請實做程式 =====

# ====================
```

<!-- grade_id: cell-d872c7c57f11c968, points: 1, grade: true, solution: false -->
```python
print('My estimation of root:', root)
```

### Try to find the zero with a tolerance of $10^{-10}$. If it works, plot the error and estimation of each step. Otherwise, state the reason why the method failed on this case.

<!-- grade_id: cell-898e5d4c01339747, points: 1, grade: true, solution: true -->
請點此cell兩下開始作答（如要打文字記得選Markdown, 寫程式則選Code, 一個cell不夠可以再新增在下方）

---

### 4. Answer the following questions under the case $c=2$.

### Plot the function to find an interval that contains the zero of $f$ if possible.

<!-- grade_id: cell-b2a7d8e20680371e, points: 0, grade: false, solution: true -->
```python
c = 2
f = g(c)

# Hint: search_range = np.arange(左端點, 右端點, 點與點之間距), 
# e.g. search_range = np.arange(0.0, 1.0, 0.01)
# ===== 請實做程式 =====

# ====================

fig, ax = plt.subplots(figsize=(9, 4))
ax.plot(search_range, f(search_range))
ax.set_title(r'$f(x)=cosh(x)+cos(x)-c$, $c=$%d' % c)
ax.grid(True)
ax.axhline(y=0)
ax.axvline(x=0)
plt.show()
```

### According to the figure above, estimate the zero of $f$.

#### For example, 

~~~python
root = 3        # 單根
root = -2, 1    # 多根
root = None     # 無解
~~~

<!-- grade_id: cell-364b7f3d6ee6a30e, points: 0, grade: false, solution: true -->
```python
# Hint: root = ?
# ===== 請實做程式 =====

# ====================
```

<!-- grade_id: cell-20fddbe6fa4c437b, points: 1, grade: true, solution: false -->
```python
print('My estimation of root:', root)

```

### Try to find the zero with a tolerance of $10^{-10}$. If it works, plot the error and estimation of each step. Otherwise, state the reason why the method failed on this case.

<!-- grade_id: cell-6721bfd92fe83ae0, points: 1, grade: true, solution: true -->
請點此cell兩下開始作答（如要打文字記得選Markdown, 寫程式則選Code, 一個cell不夠可以再新增在下方）

---

### 5. Answer the following questions under the case $c=3$.

### Plot the function to find an interval that contains the zeros of $f$ if possible.

<!-- grade_id: cell-b9e53f953d617465, points: 0, grade: false, solution: true -->
```python
c = 3
f = g(c)

# Hint: search_range = np.arange(左端點, 右端點, 點與點之間距), 
# e.g. search_range = np.arange(0.0, 1.0, 0.01)
# ===== 請實做程式 =====

# ====================

fig, ax = plt.subplots(figsize=(9, 4))
ax.plot(search_range, f(search_range))
ax.set_title(r'$f(x)=cosh(x)+cos(x)-c$, $c=$%d' % c)
ax.grid(True)
ax.axhline(y=0)
ax.axvline(x=0)
plt.show()
```

### According to the figure above, estimate the zero of $f$.

#### For example, 

~~~python
root = 3        # 單根
root = -2, 1    # 多根
root = None     # 無解
~~~

<!-- grade_id: cell-1756e3886dd01052, points: 0, grade: false, solution: true -->
```python
# Hint: root = ?
# ===== 請實做程式 =====

# ====================
```

<!-- grade_id: cell-06ec0b20844075c7, points: 1, grade: true, solution: false -->
```python
print('My estimation of root:', root)

```

### Try to find the zero with a tolerance of $10^{-10}$. If it works, plot the error and estimation of each step. Otherwise, state the reason why the method failed on this case.

<!-- grade_id: cell-676098e6b38fa075, points: 4, grade: true, solution: true -->
請點此cell兩下開始作答（如要打文字記得選Markdown, 寫程式則選Code, 一個cell不夠可以再新增在下方）

---

## Discussion

### For all cases above(c=1,2,3), do the results(e.g. error behaviors, estimations, etc) agree with the theoretical analysis?

<!-- grade_id: cell-492b11d873a7f6c2, points: 3, grade: true, solution: true -->
請點此cell兩下開始作答（如要打文字記得選Markdown, 寫程式則選Code, 一個cell不夠可以再新增在下方）